rule_files:
  - memory.yml

tests:
  - input_series:
      - series: node_memory_MemAvailable_bytes{instance="nod01.prd.grv.wma.bds.systems", job="node-exporter"}
        values: 3780547584
      - series: node_memory_MemTotal_bytes{instance="nod01.prd.grv.wma.bds.systems", job="node-exporter"}
        values: 33683968000
    alert_rule_test:
      - eval_time: 4m
        alertname: OutOfMemory
        exp_alerts: []
      - eval_time: 5m
        alertname: OutOfMemory
        exp_alerts:
          - exp_labels:
              severity: page
              instance: nod01.prd.grv.wma.bds.systems
              job: node-exporter
            exp_annotations:
              #summary: "Out of memory (instance nod01.prd.grv.wma.bds.systems)"
              #description: "Node memory is filling up (< 10% left)\n  VALUE = 65275747328\n  LABELS: {{ $labels }}"
    promql_expr_test:
      - expr: node_memory_MemTotal_bytes
        eval_time: 4m
        exp_samples:
          - labels: node_memory_MemTotal_bytes{instance="nod01.prd.grv.wma.bds.systems", job="node-exporter"}
            value: 33683968000
